---
id: 62cc5b1779e4d313466f73c5
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Nest a `link` element within the `head` element. Give it a `rel` attribute of `stylesheet` and an `href` attribute of `styles.css`.

# --hints--

Your code should have a `link` element.

```js
const link = document.querySelector('link');
assert.isNotNull(link);
```

The `link` element is a void element, it should not have an end tag `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

You should have one `link` element.

```js
const linkElementCount = document.querySelectorAll('link')?.length;
assert.strictEqual(linkElementCount, 1);
```

Your `link` element should be within your `head` element.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Your `link` element should have a `rel` attribute of `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Your `link` element should have an `href` attribute of `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
