---
id: 615f34ecc1091b4fd5a8a484
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Within your `head` element, add a `link` element with the `rel` attribute set to `stylesheet` and the `href` attribute set to `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

This will import the `Open Sans` font family, with the font weight values `400`, `700`, and `800`.

Also add a `link` element to link your `styles.css` file.

# --hints--

Your code should have two `link` elements.

```js
const linkElementCount = document.querySelectorAll('link')?.length;
assert.strictEqual(linkElementCount, 2);
```

The `link` element is a void element, it should not have an end tag `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Both of your `link` elements should have the `rel` attribute of `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.find(link => link.getAttribute('rel') === 'stylesheet'));
```

One of your `link` elements should have an `href` attribute of `styles.css`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.find(link => link.dataset?.href === 'styles.css'));
```

One of your `link` elements should have an `href` attribute of `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

```js
const url = 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800';
const links = [...document.querySelectorAll('link')];
assert(links.find(link => link.href === url));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
